---
title       : Chapter 2
description : Insert the chapter description here



--- type:NormalExercise lang:python xp:100 skills:2 key:13e65f0ce8
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}

```

*** =solution
```{python}

```

*** =sct
```{python}

```
