---
title       : One sample hypothesis testing
description : 



--- type:NormalExercise lang:python xp:100 skills:2 key:13e65f0ce8
## CHAPTER 2: Exercise 1

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 1
key: 7942f6cb9a
```

*** =instructions

Dataframe has been preloaded in `complaints`.

*** =hint

*** =pre_exercise_code
```{python}

import numpy as np
import pandas as pd
import scipy.stats as stats
import urllib.request

filename = 'https://s3.amazonaws.com/assets.datacamp.com/production/course_6777/datasets/Chicago_violent_crimes.csv.gz.csv'
urllib.request.urlretrieve(filename, "nyc_complaints_311.csv.gz")

complaints = pd.read_csv('nyc_complaints_311.csv.gz')
```

*** =sample_code
```{python}

```

*** =solution
```{python}

```

*** =sct
```{python}

```

--- type:NormalExercise lang:python xp:100 skills:2 key:9aace246e0
## CHAPTER 2: Exercise 2

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 1
key: 7942f6cb9a
```

*** =instructions

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}

```

*** =solution
```{python}

```

*** =sct
```{python}

```

